---
title: Test
description: >-
  Test


---
## daslnkdasl

```yaml
type: NormalExercise

xp: NaN

key: 9ca511ab78
```

dasdsadsa





`@solution`
```{undefined}
dasdasdas









dsadas
das
dsa
das
dad
a
```






---
## <<<New Exercise>>>

```yaml
type: PureMultipleChoiceExercise

xp: 50
skills: 2
key: 37988592a9
```




`@hint`
this is a hint





`@possible_answers`
- [1]
- 2

`@feedbacks`
sct




---
## TabEx

```yaml
type: TabExercise

xp: NaN

key: c599a73c7a
```













***



```yaml
type: NormalExercise

xp: NaN

key: 8146d74b64
```



`@instructions`
ex 1

`@hint`












***



```yaml
type: NormalExercise

xp: NaN

key: d67c3e2377
```



`@instructions`
ex 2

`@hint`












***



```yaml
type: NormalExercise

xp: NaN

key: fab9a1ed77
```



`@instructions`
ex 3

`@hint`












***



```yaml
type: NormalExercise

xp: NaN

key: e3ec2533e3
```



`@instructions`
ok seems to be rather ok!

`@hint`












---
## PureMCE

```yaml
type: PureMultipleChoiceExercise

xp: NaN

key: c855f01903
```














---
## some

```yaml
type: NormalExercise

xp: NaN

key: 5d2666dae7
```














---
## Insert exercise title here

```yaml
type: NormalExercise

xp: NaN

key: 36d434318a
```





`@pre_exercise_code`
```{undefined}
1
```
`@sample_code`
```{undefined}
3
```
`@solution`
```{undefined}
2
```
`@sct`
```{undefined}
4
```





---
## Insert exercise title here

```yaml
type: NormalExercise

xp: NaN

key: 046ca0552e
```





`@pre_exercise_code`
```{undefined}
pre
```
`@sample_code`
```{undefined}
sample
```
`@solution`
```{undefined}
solution
```
`@sct`
```{undefined}
sct
```





---
## Insert exercise title here

```yaml
type: NormalExercise

xp: NaN

key: cd7fe78246
```














---
## Insert exercise title here

```yaml
type: NormalExercise

xp: NaN

key: 81b1a735da
```





`@pre_exercise_code`
```{undefined}
pre1
```








---
## Insert exercise title here

```yaml
type: NormalExercise

xp: NaN

key: 0ad50b694f
```



`@instructions`


`@hint`




`@solution`
```{}
this is a solution and it's working
```





