---
title: Test
description: >-
  Test


---
## daslnkdasl

```yaml
type: NormalExercise

xp: 100

key: 9ca511ab78
```

dasdsadsa





`@solution`
```{python}
dasdasdas









dsadas
das
dsa
das
dad
a
```






---
## <<<New Exercise>>>

```yaml
type: PureMultipleChoiceExercise

xp: 50
skills: 2
key: 37988592a9
```




`@hint`
this is a hint





`@possible_answers`
- [1]
- 2

`@feedbacks`
sct




---
## TabEx

```yaml
type: TabExercise

xp: 100

key: c599a73c7a
```













***



```yaml
type: NormalExercise

xp: 25

key: 8146d74b64
```



`@instructions`
ex 1

`@hint`












***



```yaml
type: NormalExercise

xp: 25

key: d67c3e2377
```



`@instructions`
ex 2

`@hint`












***



```yaml
type: NormalExercise

xp: 25

key: fab9a1ed77
```



`@instructions`
ex 3

`@hint`












***



```yaml
type: NormalExercise

xp: 25

key: e3ec2533e3
```



`@instructions`
ok seems to be rather ok!

`@hint`












---
## PureMCE

```yaml
type: PureMultipleChoiceExercise

xp: 50

key: c855f01903
```














---
## some

```yaml
type: NormalExercise

xp: 100

key: 5d2666dae7
```














---
## Insert exercise title here

```yaml
type: NormalExercise

xp: 100

key: 36d434318a
```





`@pre_exercise_code`
```{python}
1
```
`@sample_code`
```{python}
3
```
`@solution`
```{python}
2
```
`@sct`
```{python}
4
```





---
## Insert exercise title here

```yaml
type: NormalExercise

xp: 100

key: 046ca0552e
```





`@pre_exercise_code`
```{python}
pre
```
`@sample_code`
```{python}
sample
```
`@solution`
```{python}
solution
```
`@sct`
```{python}
sct
```





---
## Insert exercise title here

```yaml
type: NormalExercise

xp: 100

key: cd7fe78246
```














---
## Insert exercise title here

```yaml
type: NormalExercise

xp: 100

key: 81b1a735da
```





`@pre_exercise_code`
```{python}
pre1
```








---
## Insert exercise title here

```yaml
type: NormalExercise

xp: 100

key: 0ad50b694f
```







`@solution`
```{python}
this is a solution and it's working
```






---
## Insert exercise title here

```yaml
type: VideoExercise

xp: 50

key: 6882d46654
```

`@projector_key`
2d71e993128ceca527df673393ba092b